---
title: Introduction
description: 'Dans ce chapitre, nous allons mettre l''accent sur l''utilisation de R comme outils d''analyse des données qualitatives'
---

## Exercice 1

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Exercice n° 1.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Exercie intro

```yaml
type: NormalExercise
key: 75c4b6c136
xp: 100
```

Notre jeux des données utilise la base `auto2.txt` les variables sont:

1. les variables manifestes de formation
2. les variables de dfdkljmklfhf
3. dsgfdhfh

`@instructions`
1. utiliser la library necessaire à l'utilisataodjnbgf
2. djhfdkjglmfdkhfg
3. la library `dfgfdhgfh` rets ksdgmkjkmdfh

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
#données
library(lavaan)

```

`@solution`
```{r}
#données
library(lavaan)

```

`@sct`
```{r}

```

---

## Introduction

```yaml
type: VideoExercise
key: ef9e2fcbd7
xp: 50
```

`@projector_key`
813e6db35a04a1c8ecd4454fde5f81c1
